## UBS AM Python Self Assessment
### Welcome to the Python Self Assessment designed as a benchmark of skills required as a prerequisite for the UBS/Stevens Foundations in Quant and Data Science Certificate.  Please feel free to take this assessment as many times as you want.  We recommend you first take a Python MOOC to study the requisite skill.

#### Problem 1:  We'll start with the basics, please import the pandas, numpy, and sklearn.  These are pre-installed in the conda environment, so all you have to do is import them to this notebook. (1 pt)

<!-- grade_id: cell-2da0ecf3f975203b, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

### END SOLUTION
```

<!-- grade_id: cell-0b55666ce32d716c, points: 1, grade: true, solution: false -->
```python
# Problem 1 Test Cell
### BEGIN HIDDEN TESTS
assert pandas.__version__.__contains__('0.2')
assert numpy.__version__.__contains__('1.1')
assert sklearn.__version__.__contains__('0.2')
### END HIDDEN TESTS
```

#### Problem 2: Create a list called 'foods' that contains the following items: apples, oranges, beer, pasta.  Create a second list called 'prices'.  Apples cost 3, oranges cost 2, beer costs 4.5, and pasta costs 10. Finally, create a list called 'categories' that represent your foods list categories.  Apples and oranges are fruit, beer is alcohol, and pasta is in the grains category. (1 pt)

<!-- grade_id: cell-9f4d01abd6b13157, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

### END SOLUTION
```

<!-- grade_id: cell-9b800b425b93748d, points: 1, grade: true, solution: false -->
```python
# Problem 2 Test Cell
### BEGIN HIDDEN TESTS
assert foods == ['apples','oranges','beer','pasta']
assert prices == [3,2,4.5,10]
assert categories == ['fruit','fruit','alcohol','grains']
### END HIDDEN TESTS
```

#### Problem 3: Build a pandas dataframe called 'groceries_df' that contains each of your lists as the columns of the dataframe. Use the list names as the column names.  The expected shape of the dataframe is 4 rows and 3 columns. (2pts)

<!-- grade_id: cell-72e1f4f594e73ddc, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

### END SOLUTION
```

<!-- grade_id: cell-c38250956ce785fa, points: 2, grade: true, solution: false -->
```python
# Problem 3 Test Cell
### BEGIN HIDDEN TESTS
assert type(groceries_df) == pandas.core.frame.DataFrame
assert groceries_df.shape == (4,3)
assert list(groceries_df.columns) == ['foods','prices','categories']
### END HIDDEN TESTS
```

#### Problem 4: Calculate the average price of the fruits category in your data frame.  First create a groupby object called cat_grouped.  Then use apply, transform, or agg to calculate the average price of fruit and call that fruit_ave_price. (2 pts)

<!-- grade_id: cell-8c14883c6f89b34c, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

### END SOLUTION
```

<!-- grade_id: cell-835866a4d0c0c630, points: 2, grade: true, solution: false -->
```python
# Problem 4 Test Cell
### BEGIN HIDDEN TESTS
assert type(cat_grouped) == pandas.core.groupby.groupby.SeriesGroupBy
assert fruit_ave_price == 2.5
### END HIDDEN TESTS
```

#### Problem 5: add a new column to the groceries_df data frame indicating the aisle in which you can find the food in the store using the list given to you below.  Call the new column 'aisle'. (2 pts)

<!-- grade_id: cell-c7bf11e2b0e578dd, points: 0, grade: false, solution: true -->
```python
aisle_list = [2,7,numpy.nan,3]

### BEGIN SOLUTION

### END SOLUTION
```

<!-- grade_id: cell-4e3784dc93216da4, points: 2, grade: true, solution: false -->
```python
# Problem 5 Test Cell
### BEGIN HIDDEN TESTS
assert list(groceries_df.columns) == ['foods','prices','categories','aisle']
assert groceries_df['aisle'].equals(pandas.Series(aisle_list))
### END HIDDEN TESTS
```

#### Problem 6: Create a new dataframe called groc_df_new that filters out any null values from the groceries_df dataframe. (2 pts)

<!-- grade_id: cell-11c82915c5280fde, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

### END SOLUTION
```

<!-- grade_id: cell-1b19d6ca46f62443, points: 2, grade: true, solution: false -->
```python
# Problem 6 Test Cell
### BEGIN HIDDEN TESTS
assert groc_df_new.shape == (3,4)
assert groc_df_new['foods'].equals(pandas.Series(['apples','oranges','pasta'],index=[0,1,3]))
### END HIDDEN TESTS
```

#### Problem 7: Write a python function that returns the square of it's input, call this function 'squares.'  Then create a new column by applying it to the 'aisle' column of groc_df_new.  Call the new column 'aisle_squared.' (3 pts)

<!-- grade_id: cell-19d3892777a9617a, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

### END SOLUTION
```

Output:
```
C:\Users\bacro\Anaconda3\lib\site-packages\ipykernel_launcher.py:2: SettingWithCopyWarning: 
A value is trying to be set on a copy of a slice from a DataFrame.
Try using .loc[row_indexer,col_indexer] = value instead

See the caveats in the documentation: http://pandas.pydata.org/pandas-docs/stable/indexing.html#indexing-view-versus-copy
```

<!-- grade_id: cell-727433c9cd5bf5f7, points: 3, grade: true, solution: false -->
```python
# Problem 7 Test Cell
### BEGIN HIDDEN TESTS
assert squares(3) == 9
assert groc_df_new['aisle_squared'].equals(pandas.Series([4.0,49.0,9.0],index=[0,1,3]))
### END HIDDEN TESTS
```

#### Problem 8: Load two dataframes, called df_A and df_B, from the below paths and label the columns with the lists provided below.  Next merge the two data frames using date and ticker as the common keys, call this df_C.  Lastly, calculate the daily series of net dollar exposures (buy - sell), call this net_dollars. (5 pts)

<!-- grade_id: cell-ce6ce49d8d59eeca, points: 0, grade: false, solution: true -->
```python
df_A_path = './data/df_A.csv'
df_B_path = './data/df_B.csv'

### BEGIN SOLUTION

### END SOLUTION
```

Output:
```
-9528.899999999794
```

<!-- grade_id: cell-e30866559291653f, points: 5, grade: true, solution: false -->
```python
# Problem 8 Test Cell
### BEGIN HIDDEN TESTS
from math import isclose
assert df_A.shape == (8,3)
assert df_B.shape == (8,4)
assert df_C.shape == (8,6)
assert isclose(net_dollars.dollars.sum(),-9528,abs_tol=1)
### END HIDDEN TESTS
```

#### Problem 9: Load the Boston data set from the R MASS package into a pandas dataframe called 'boston'.  Create new columns by transforming the columns 'rm' (rooms) and 'medv' (median values) into z-scores. Call the new columns 'z_rm' and 'z_medv.' (5 pts)

<!-- grade_id: cell-b52b96fbbcf99df6, points: 0, grade: false, solution: true -->
```python
boston_path = './data/boston.csv'

### BEGIN SOLUTION

### END SOLUTION
```

<!-- grade_id: cell-d1c3f86a8e0de558, points: 5, grade: true, solution: false -->
```python
# Problem 9 Test Cell
### BEGIN HIDDEN TESTS
assert type(boston) == pandas.core.frame.DataFrame
assert boston.shape == (506,17)
assert list(boston.columns) == ['Unnamed: 0','crim','zn','indus', 'chas', 'nox', 'rm', 'age', 'dis', 'rad', 'tax',
                                 'ptratio', 'black', 'lstat', 'medv', 'z_rm', 'z_medv']

z_rm_test = (boston['rm'] - boston['rm'].mean())/boston['rm'].std(ddof=0)
z_medv_test = (boston['medv'] - boston['medv'].mean())/boston['medv'].std(ddof=0)
assert boston.z_rm.equals(z_rm_test)
assert boston.z_medv.equals(z_medv_test)
### END HIDDEN TESTS
```

#### Problem 10: Use sci-kit learn's linear regression model to build a model to predict home values(z_medv) from the number of rooms in the house (z_rm).  Fit the model, then label the coefficient value as 'beta', the intercept value as 'alpha', and the R-squared value as 'rsq'. (7pts)

<!-- grade_id: cell-9ead323bde633bde, points: 0, grade: false, solution: true -->
```python
from sklearn.linear_model import LinearRegression

### BEGIN SOLUTION

### END SOLUTION
```

<!-- grade_id: cell-ad802e64d32e9a98, points: 7, grade: true, solution: false -->
```python
# Problem 10 Test Cell
### BEGIN HIDDEN TESTS
assert isclose(alpha,0,abs_tol=0.01)
assert isclose(beta,0.69,abs_tol=0.01)
assert isclose(rsq,0.48,abs_tol=0.01)
### END HIDDEN TESTS
```